---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Ex 1.1

```yaml
type: NormalExercise 
lang: sql
xp: 100 
skills: 1
key: b6305f0655   
```


Do some data science.


`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{sql}

```


`@solution`

```{sql}

```


`@sct`

```{python}

```


---

## Intro to T-SQL / Microsoft SQL

```yaml
type: VideoExercise 
xp: 50 
key: 0851e37bc2   
```

`@projector_key`
bf2ab310c48319c4cb626ccd81da9fe6
